---
title: 'Operating on Missing Data'
description: 'After a quick understanding of what are the different types of missing data, you would deal with in real-life problems, let''s move further by learning how would you operate on different types of missing data.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Forward and Backward Filling

```yaml
type: VideoExercise
key: 8a0b5f7636
xp: 50
```

`@projector_key`
0c5f8c88ecf599b36dbbe2e97361d32c
